# Lab 2 Main: Data Manipulation

Apply data wrangling skills to two datasets: `WeatherTrips.csv` and `products.csv`.

Complete all questions and run the submission cells at the end.

```python
import pandas as pd
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab-02-main.ipynb")
```

---
## Part 1: Warmup

Complete these warmup questions to practice basic Pandas operations using the `WeatherTrips.csv` dataset, which contains daily bike trip counts and weather data.

**Dataset Columns:**
- `DATE` – Date of the observation
- `NumTrips` – Number of bike trips that day
- `PRCP` – Precipitation (inches)
- `AWND` – Average wind speed
- `TAVG` – Average temperature (°F)

### Question 1: Loading
Load the `WeatherTrips.csv` file into a DataFrame assigned to the variable `df`.

```python
...
```

```python
grader.check("q1")
```

### Question 2: Filtering
Create a new DataFrame called `warm_days` that only contains days where the average temperature (`TAVG`) is **50 degrees or higher**.

```python
...
```

```python
grader.check("q2")
```

### Question 3: Sorting
Assign the 5 days with the **highest** number of trips (`NumTrips`) to a new DataFrame called `busiest_days`.

```python
...
```

```python
grader.check("q3")
```

### Question 4: Grouping & Aggregation
First, create a new column called `MONTH` by extracting the first 7 characters from the `DATE` column (e.g., "2017-01").

Then, calculate the **average number of trips** for each month. Assign the result (a Series) to `monthly_avg`.

*Hint: To extract the first 7 characters from the `DATE` column, use `df["MONTH"] = df["DATE"].str[:7]`*

```python
...
```

```python
grader.check("q4")
```

### Question 5: Pivot Table
Create a Pivot Table that shows the **total number of trips** for each month:
- Rows: `MONTH`
- Values: `NumTrips` (sum)

Assign this table to `monthly_totals`.

*Hint: Use `pivot_table()` with `aggfunc='sum'`*

```python
...
```

```python
grader.check("q5")
```

---
## Part 2: Main Questions

Now apply the same skills to a new dataset: `products.csv`, which contains information about grocery store products.

**Dataset Columns:**
- `product_id` – Unique product identifier
- `product_name` – Name of the product
- `aisle_id` – Aisle identifier (which aisle the product is in)
- `department_id` – Department identifier (which department the product belongs to)

### Question 6: Loading
Load the `products.csv` file into a DataFrame assigned to the variable `products`.

```python
...
```

```python
grader.check("q6")
```

### Question 7: Filtering
Create a new DataFrame called `dept_7` that only contains products from **department 7**.

```python
...
```

```python
grader.check("q7")
```

### Question 8: Sorting
Sort the DataFrame by `product_name` in **alphabetical order** (ascending). Assign the first 5 rows to a new DataFrame called `first_5_products`.

```python
...
```

```python
grader.check("q8")
```

### Question 9: Grouping & Aggregation
Count the **number of products** in each department (`department_id`). Assign the result (a Series) to `products_per_dept`.

```python
...
```

```python
grader.check("q9")
```

### Question 10: Pivot Table
Create a Pivot Table that shows the **count** of products for each department:
- Rows: `department_id`
- Values: `product_id` (count)

Assign this table to `dept_totals`.

*Hint: Use `pivot_table()` with `aggfunc='count'`*

```python
...
```

```python
grader.check("q10")
```

---
## Submission

Run the cells below to generate your submission ZIP file.

```python
!jupyter nbconvert --to script lab-02-main.ipynb
```

```python
grader.export(pdf=False, force_save=True, run_tests=True, files=['lab-02-main.py'])
```

Tests for q1 (points: 1)
```
>>> assert isinstance(df, pd.DataFrame)
>>> assert len(df) == 90
```


Tests for q10 (points: 1)
```
>>> assert 1 in dept_totals.index
>>> assert dept_totals['product_id'].sum() == len(products)
```


Tests for q2 (points: 1)
```
>>> assert len(warm_days) < len(df)
>>> assert warm_days['TAVG'].min() >= 50
```


Tests for q3 (points: 1)
```
>>> assert len(busiest_days) == 5
>>> assert busiest_days.iloc[0]['NumTrips'] == df['NumTrips'].max()
```


Tests for q4 (points: 1)
```
>>> assert '2017-01' in monthly_avg.index
>>> assert len(monthly_avg) == 3
>>> assert monthly_avg['2017-01'] < monthly_avg['2017-03']
```


Tests for q5 (points: 1)
```
>>> assert '2017-01' in monthly_totals.index
>>> assert len(monthly_totals) == 3
>>> assert monthly_totals['NumTrips'].sum() == df['NumTrips'].sum()
```


Tests for q6 (points: 1)
```
>>> assert isinstance(products, pd.DataFrame)
>>> assert len(products) == 49688
```


Tests for q7 (points: 1)
```
>>> assert len(dept_7) < len(products)
>>> assert dept_7['department_id'].unique().tolist() == [7]
```


Tests for q8 (points: 1)
```
>>> assert len(first_5_products) == 5
>>> assert first_5_products.iloc[0]['product_name'] == products['product_name'].min()
```


Tests for q9 (points: 1)
```
>>> assert 1 in products_per_dept.index
>>> assert products_per_dept.sum() == len(products)
```
