# Lab 2 Warmup: Data Manipulation

Practice data wrangling with the `WeatherTrips.csv` dataset.

**Dataset:** Daily bike trip counts and weather data
- `DATE`, `NumTrips`, `PRCP`, `AWND`, `TAVG`

```python
import pandas as pd
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab-02-warmup.ipynb")
```

### Question 1: Loading
Load the `WeatherTrips.csv` file into a DataFrame assigned to the variable `df`.

```python
...
```

```python
grader.check("q1")
```

### Question 2: Filtering
Create a new DataFrame called `warm_days` that only contains days where the average temperature (`TAVG`) is **50 degrees or higher**.

```python
...
```

```python
grader.check("q2")
```

### Question 3: Sorting
Assign the 5 days with the **highest** number of trips (`NumTrips`) to a new DataFrame called `busiest_days`.

```python
...
```

```python
grader.check("q3")
```

### Question 4: Grouping & Aggregation
First, create a new column called `MONTH` by extracting the first 7 characters from the `DATE` column (e.g., "2017-01").

Then, calculate the **average number of trips** for each month. Assign the result (a Series) to `monthly_avg`.

*Hint: To extract the first 7 characters from the `DATE` column, use `df["MONTH"] = df["DATE"].str[:7]`*

```python
...
```

```python
grader.check("q4")
```

### Question 5: Pivot Table
Create a Pivot Table that shows the **total number of trips** for each month:
- Rows: `MONTH`
- Values: `NumTrips` (sum)

Assign this table to `monthly_totals`.

*Hint: Use `pivot_table()` with `aggfunc='sum'`*

```python
...
```

```python
grader.check("q5")
```

---
## Submission

Run the cells below to generate your submission ZIP file.

```python
!jupyter nbconvert --to script lab-02-warmup.ipynb
```

```python
grader.export(pdf=False, force_save=True, run_tests=True, files=['lab-02-warmup.py'])
```

Tests for q1 (points: 1)
```
>>> assert isinstance(df, pd.DataFrame)
>>> assert len(df) == 90
```


Tests for q2 (points: 1)
```
>>> assert len(warm_days) < len(df)
>>> assert warm_days['TAVG'].min() >= 50
```


Tests for q3 (points: 1)
```
>>> assert len(busiest_days) == 5
>>> assert busiest_days.iloc[0]['NumTrips'] == df['NumTrips'].max()
```


Tests for q4 (points: 1)
```
>>> assert '2017-01' in monthly_avg.index
>>> assert len(monthly_avg) == 3
>>> assert monthly_avg['2017-01'] < monthly_avg['2017-03']
```


Tests for q5 (points: 1)
```
>>> assert '2017-01' in monthly_totals.index
>>> assert len(monthly_totals) == 3
>>> assert monthly_totals['NumTrips'].sum() == df['NumTrips'].sum()
```
